Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

<!-- grade_id: set_a, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

<!-- grade_id: foo, points: 1, grade: true, solution: false -->
```python
print("Success!")
```

<!-- grade_id: bar, points: 1, grade: true, solution: false -->
```python
assert a == 1
```

<!-- grade_id: baz, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: quux, points: 3, grade: true, solution: true -->
```python
# your code here
raise NotImplementedError
```

```python
print("Don't change this cell!")
```

This cell shouldn't be changed.